# Homework 1 - 
**Problem 1**
If statement:

Given an integer n, perform the following conditional actions:

    If n is odd, set output to "Hello World"
    If n is even and in the inclusive range of 2 to 5, set output to "Goodbye World"
    If n is even and in the inclusive range of 6 to 20, set output to "Hello World"
    If n is even and greater than 20, set output to "Goodbye World"

Example:

input: 
    
    n = 2

output:

    output = "Goodbye World"

<!-- grade_id: cell-8ce4028ecf5e6219, points: 0, grade: false, solution: true -->
```python
n = 3

def problem1(n):
    output = 0
    
    if n % 2 != 0: # if n is odd
        output = 'Hello World'
    elif n >= 2 and n <= 5:
        output = 'Goodbye World'
    elif n >= 6 and n <= 20:
        output = 'Hello World'
    else:
        output = 'Goodbye World'
    
    return(output)
```

<!-- grade_id: cell-265d8840d2710be9, points: 25, grade: true, solution: false -->
```python
# Test cases for Problem Set 1
assert problem1(3) == "Hello World"
assert problem1(2) == "Goodbye World"
assert problem1(18) == "Hello World"
```

<!-- grade_id: cell-fa5f8649a0295e07, points: 25, grade: true, solution: false -->
```python
assert problem1(29) == "Hello World"
assert problem1(24) == "Goodbye World"
```

**Problem 2**
(For-loop, basic list manipulation)

You and $K-1$ friends want to buy $N$ different pieces of candy.  Candy piece $i$ costs $p_i$. The candy shop owner is greedy and wants to maximize his number of new customers, so he increases the sale price of the candy for repeated customers; more precisely, if a customer has already purchased $x$ pieces of candy, then the cost of candy $i$ is $C_i*(x+1)$.

Find and print the minimum cost for your group to purchase all $N$ pieces of candy.

The input:

    N: the number of candies to buy
    K: the number of people
    C: the cost of candy

The output:

    the total cost

**Hint:** You may want to use the list comprehension function.


Example:

Input:

    N: 4
    K: 2
    C: [1,3,4,5]

Output:

    cost = 4 + 5 + 3 * 2 + 1 * 2 = 17

(Why 17 is the minimum cost in this case? You would always like your group to purchase the most expensive candies first. In this way, you will get the multipliers on candy costs for the cheap candies. In this case, two people in your group first purchase candies with costs 4 and 5. And then they will purchase candies with costs 3 and 1. When purchasing candies with costs 3 and 1, both of persons in your group are repeated customers, so they pay 3\*2 and 1\*2.)


<!-- grade_id: cell-5ad1d968ba47b22c, points: 0, grade: false, solution: true -->
```python
#Uncomment to fix the test case
import random
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]

def problem2(N, K, C):
    #sort costs decreasing
    C = sorted(C,reverse = True)
    
    #variable to store the cost
    cost = 0
    
    if N % K == 0:
        num = N/K
    else:
        num = int(N/K) + 1
    
    rep_num = list(range(1, int(num+1)))
    rep_num = rep_num*K
    rep_num = sorted(rep_num)

    for i in range(len(C)):
        cost += C[i]*rep_num[i]

    return(cost)

```

<!-- grade_id: cell-b6244c846cd441c4, points: 25, grade: true, solution: false -->
```python
random.seed(100)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 8838
```

<!-- grade_id: cell-062e4c6c399bd36f, points: 25, grade: true, solution: false -->
```python
random.seed(10)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 193878
```